rule_files:
  - trustyai-alerting.rules.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      - series: probe_success:burnrate5m{instance="trustyai-service-operator-controller-manager"}
        values: "0x60"
      - series: probe_success:burnrate1h{instance="trustyai-service-operator-controller-manager"}
        values: "0x60"
      - series: probe_success:burnrate30m{instance="trustyai-service-operator-controller-manager"}
        values: "0x60"
      - series: probe_success:burnrate6h{instance="trustyai-service-operator-controller-manager"}
        values: "0x60"
      - series: probe_success:burnrate2h{instance="trustyai-service-operator-controller-manager"}
        values: "0x60"
      - series: probe_success:burnrate1d{instance="trustyai-service-operator-controller-manager"}
        values: "0x60"
    alert_rule_test:
      - eval_time: 1h
        alertname: TrustyAI Controller Probe Success 5m and 1h Burn Rate high
        exp_alerts: []
      - eval_time: 1h
        alertname: TrustyAI Controller Probe Success 30m and 6h Burn Rate high
        exp_alerts: []
      - eval_time: 1h
        alertname: TrustyAI Controller Probe Success 2h and 1d Burn Rate high
        exp_alerts: []

  - interval: 1m
    input_series:
      - series: probe_success:burnrate5m{instance="trustyai-service-operator-controller-manager"}
        values: "1+1x60"
      - series: probe_success:burnrate1h{instance="trustyai-service-operator-controller-manager"}
        values: "1+1x60"
    alert_rule_test:
      - eval_time: 2m
        alertname: TrustyAI Controller Probe Success 5m and 1h Burn Rate high
        exp_alerts:
          - exp_labels:
              alertname: TrustyAI Controller Probe Success 5m and 1h Burn Rate high
              instance: "trustyai-service-operator-controller-manager"
              severity: critical
            exp_annotations:
              summary: "TrustyAI Controller Probe Success 5m and 1h Burn Rate high"
              message: "High error budget burn for trustyai-service-operator-controller-manager (current value: 3)."
              triage: 'https://gitlab.cee.redhat.com/service/managed-tenants-sops/-/blob/main/RHODS/Model-Serving/rhoai-trustyai-controller-probe-success-burn-rate.md'

  - interval: 1m
    input_series:
      - series: probe_success:burnrate30m{instance="trustyai-service-operator-controller-manager"}
        values: "1+1x60"
      - series: probe_success:burnrate6h{instance="trustyai-service-operator-controller-manager"}
        values: "1+1x60"
    alert_rule_test:
      - eval_time: 15m
        alertname: TrustyAI Controller Probe Success 30m and 6h Burn Rate high
        exp_alerts:
          - exp_labels:
              alertname: TrustyAI Controller Probe Success 30m and 6h Burn Rate high
              instance: "trustyai-service-operator-controller-manager"
              severity: critical
            exp_annotations:
              summary: "TrustyAI Controller Probe Success 30m and 6h Burn Rate high"
              message: "High error budget burn for trustyai-service-operator-controller-manager (current value: 16)."
              triage: 'https://gitlab.cee.redhat.com/service/managed-tenants-sops/-/blob/main/RHODS/Model-Serving/rhoai-trustyai-controller-probe-success-burn-rate.md'

  - interval: 1m
    input_series:
      - series: probe_success:burnrate2h{instance="trustyai-service-operator-controller-manager"}
        values: "1+1x60"
      - series: probe_success:burnrate1d{instance="trustyai-service-operator-controller-manager"}
        values: "1+1x60"
    alert_rule_test:
      - eval_time: 1h
        alertname: TrustyAI Controller Probe Success 2h and 1d Burn Rate high
        exp_alerts:
          - exp_labels:
              alertname: TrustyAI Controller Probe Success 2h and 1d Burn Rate high
              instance: "trustyai-service-operator-controller-manager"
              severity: warning
            exp_annotations:
              summary: "TrustyAI Controller Probe Success 2h and 1d Burn Rate high"
              message: "High error budget burn for trustyai-service-operator-controller-manager (current value: 61)."
              triage: 'https://gitlab.cee.redhat.com/service/managed-tenants-sops/-/blob/main/RHODS/Model-Serving/rhoai-trustyai-controller-probe-success-burn-rate.md'
